<a href="https://colab.research.google.com/github/VishalRavichandran27/hds5210-2023/blob/main/week06/week06_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE=json.load(open('allergies.json'))
type(ALLERGIES_FILE)
```

Output:
```
dict
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_count(ALLERGIES_FILE):
  """(dict)->int
Here i have created a variable as events and i have run the for loop so that i
loop over the entries and i have put the if condition to access the certainty
from the entry so if the certainty is confirmed it adds up the events so that
we get to see how many allergies are recorded.

>>> allergy_count(ALLERGIES_FILE)
4

"""
  events=0
  for i in ALLERGIES_FILE['entry']:
    if i['resource']['reaction'][0]['certainty']=='confirmed':
      events+=1
  return events
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

```python
import doctest
doctest.run_docstring_examples(allergy_count,globals(),verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_count(ALLERGIES_FILE)
Expecting:
    4
ok
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
# Put your solution here
def patient_count(ALLERGIES_FILE):
  """
  Here i have created a empty listas patients and i have called the entry from the jason file
  to access the file and then i called the resource, patient,display and attached in
  the empty list called as patientand the set that and counted the length of the patient list.
  (dict)->int
  >>> patient_count(ALLERGIES_FILE)
  2
  """
  patients=[]
  for i in ALLERGIES_FILE['entry']:
    patients.append(i['resource']['patient']['display'])
  patients=set(patients)
  return len(patients)
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

```python
import doctest
doctest.run_docstring_examples(patient_count,globals(),verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    patient_count(ALLERGIES_FILE)
Expecting:
    2
ok
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
# Put your solution here
def allergy_per_patient(ALLERGIES_FILE):
  """
  Here i have created the emty list as number allergies of patients and then called the
  entry from the jason file and form that i have called the recourse,patient,display,and
  attached list and i have created a dictionary later are we are counting the number of allergies of patients
  and then the patients reactions.
  (dict)->dict
  >>> allergy_per_patient(ALLERGIES_FILE)
  {'Jason Argonaut': 3, 'Paul Boal': 1}
  """
  number_allergies_of_patients=[]
  for i in ALLERGIES_FILE['entry']:
    number_allergies_of_patients.append(i['resource']['patient']['display'])
  patients_reactions={}
  for i in number_allergies_of_patients:
    count=number_allergies_of_patients.count(i)
    patients_reactions[i]=count
  return patients_reactions
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

```python
import doctest
doctest.run_docstring_examples(allergy_per_patient,globals(),verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_per_patient(ALLERGIES_FILE)
Expecting:
    {'Jason Argonaut': 3, 'Paul Boal': 1}
ok
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
from typing import List
# Put your solution here
def allergy_list(ALLERGIES_FILE):
  """
  Here we are opening the empty listas patients_history and calling the entry and
  in that i have called the recource,patients,display,text,reactions,manifestation,text in three variable called patients,substance,reaction
  and i have atteched it in the patients history list.
  (dict)->list
  >>> allergy_list(ALLERGIES_FILE)
  [['Jason Argonaut', 'PENICILLIN G', 'Hives'], ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'], ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'], ['Paul Boal', 'PENICILLIN G', 'Bruising']]
  """

  patients_history=[]
  for i in ALLERGIES_FILE['entry']:
    patients=i['resource']['patient']['display']
    substance=i['resource']['substance']['text']
    reaction=i['resource']['reaction'][0]['manifestation'][0]['text']
    patients_history.append([patients,substance,reaction])
  return patients_history
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

```python
import doctest
doctest.run_docstring_examples(allergy_list,globals(),verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_list(ALLERGIES_FILE)
Expecting:
    [['Jason Argonaut', 'PENICILLIN G', 'Hives'], ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'], ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'], ['Paul Boal', 'PENICILLIN G', 'Bruising']]
ok
```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_reaction(ALLERGIES_FILE,name='patient',item='drug'):
  """
  Here i have i have opened a empty string as report
  then to access the entry i have run the for loop and as similar to the before
  i have called the resource,patient,display,text,substance,manifestation,text in the three variables
  as patients,substance,reaction and then i kept patients equal to name,substance equal to item
  and then got the report.
  >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
  'Hives'

  """
  report=''
  for i in ALLERGIES_FILE['entry']:
      patients=i['resource']['patient']['display']
      substance=i['resource']['substance']['text']
      reaction=i['resource']['reaction'][0]['manifestation'][0]['text']
      if patients==name and substance==item:
        report=reaction
  if report=='':
      report=None
  return report
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

```python
import doctest
doctest.run_docstring_examples(allergy_reaction,globals(),verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
Expecting:
    'Hives'
ok
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

